# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<=12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd

def proportion_of_education():
   df=pd.read_csv('assets/NISPUF17.csv', index_col=0)
   columns_to_keep=df['EDUC1']
   sortedEDUC1=df.sort_values(by=['EDUC1'])
   t=columns_to_keep.count()
   s=columns_to_keep.value_counts()    #series 
   proportion= s/t
   prop=proportion.round(2)
   # code for proportion of 1
   prop1= {"less than high school": prop[1],"high school":prop[2],"more than high school but not college":prop[3],"college":prop[4]}
   return prop1  
   #prop1=dict({"less than high school": prop[1],"high school":prop[2],"more than high school but not college":prop[3],"college":prop[4]})
   #the code line above isn't right because this isn't the right way to create a dictionary
   raise NotImplementedError()
print(proportion_of_education()) 
```

Output:
```
{'less than high school': 0.1, 'high school': 0.17, 'more than high school but not college': 0.25, 'college': 0.48}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
import pandas as pd

'''we want to compare the avg number of vaccines in relation to the children who have received breast milk'''
# CBF_01 -> 1 : children have received breast milk
#P_NUMFLU  the total number of seasonal influenza doses for each child
df = pd.read_csv("assets/NISPUF17.csv", index_col=0)    #df.head()
sortedCBF=df.sort_values(by=['CBF_01','P_NUMFLU']) #we sort the initial DataFrame (df) according to the CBF_01 and P_NUMFLU
#sortedCBF
both= sortedCBF.loc[:,['CBF_01','P_NUMFLU']].dropna()  #Create a copy Dataframe with only these 2 columns and drop all the NA from 
                                                       # the data (the second column has NA values whilst the first contains the values 
                                                       # 1:YES has received BM /2:No/77:dont know/99:Missing)
#both
#get_cbf1 = both.where(both['CBF_01']==1).dropna()               #We get all the CBF_01 column where the value is ==1
#get_Flu1 = get_cbf1.where(get_cbf1['P_NUMFLU']!=0).dropna()      #We get all the P_NUMFLU column where the value is >=1 and !=NA
#get_Flu1
#totalDoses1 = get_Flu1['P_NUMFLU'].sum() /len(get_Flu1)    #calculation of mean average of the second column (P_NUMFLU)
#totalDoses1.round(1)                                     #the average number of vaccines in relation to the children who 
                                                        #have received breast milk

'''and the avg number of vaccines for children who haven't received breast milk '''
#CBF_01 -> 2: children who haven't received breast milk
#P_NUMFLU  the total number of seasonal influenza doses for each child
get_cbf2 = both.where(both['CBF_01']==1).dropna()               #We get all the CBF_01 column where the value is ==2
get_Flu2 = get_cbf2.where(get_cbf2['P_NUMFLU']==0).dropna()      #We get all the P_NUMFLU column where the value is >=1 and !=NA
get_Flu2
#totalDoses2 = get_Flu2['P_NUMFLU'].sum() /len(get_Flu2)    #calculation of mean average of the second column (P_NUMFLU)
#totalDoses2  

```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
import pandas as pd
def average_influenza_doses():
    df = pd.read_csv("assets/NISPUF17.csv", index_col=0)    #df.head()
    sortedCBF=df.sort_values(by=['CBF_01','P_NUMFLU'])
    both= sortedCBF.loc[:,['CBF_01','P_NUMFLU']]
    both
    get_cbf1 = both.where(both['CBF_01']==1)  #get_cbf1 gets -> CBF_01 : everything shorted  P_NUMFLU : Everything shorted
    get_Flu1 = get_cbf1.where(get_cbf1['P_NUMFLU']>=0).mean()
    #print(get_Flu1['P_NUMFLU'])

    #----------------------------------------------------------------------------------------
    get_cbf2 = both.where(both['CBF_01']==2)  #get_cbf1 gets -> CBF_01 : everything shorted  P_NUMFLU : Everything shorted
    #get_cbf1
    get_Flu2 = get_cbf2.where(get_cbf2['P_NUMFLU']>=0).mean()
    return (get_Flu1['P_NUMFLU'],get_Flu2['P_NUMFLU'])
    raise NotImplementedError()
print(average_influenza_doses())
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
import pandas as pd
def chickenpox_by_sex():
    df=pd.read_csv('assets/NISPUF17.csv',index_col=0)
    df
    #SEQNUMC children  SEX of a child  HAD_CPOX – did child ever have chicken pox   P_NUMVRC – total number of varicella doses
    sortSEX=df.sort_values(by=['SEX','HAD_CPOX'])
    both= sortSEX.loc[:,['SEX', 'HAD_CPOX', 'P_NUMVRC']]
    #both
    #--------------------------------------------------------------------------------
    fem_cpox = both.where(both['SEX']==2)  #ALL Female Children (SEX ==2) 
    p_numvrc_F= fem_cpox.where(fem_cpox['P_NUMVRC']>=1)   #...WHO WERE VACCINATED FOR CPOX

    had_cpox1= p_numvrc_F.where(p_numvrc_F['HAD_CPOX']==1)   # 1 : is for those who HAD CPOX
    had_cpox2= p_numvrc_F.where(p_numvrc_F['HAD_CPOX']==2)   # 2 : is for those who DIDN'T HAVE CPOX

    resultF= had_cpox1['SEX'].count()/had_cpox2['SEX'].count()   #Here we ensure that we take only the girls
    #print(resultF)
    #----------------------------------------------------------------------------------
    m_cpox = both.where(both['SEX']==1)  #ALL Male Children (SEX ==1) 
    p_numvrc_M= m_cpox.where(m_cpox['P_NUMVRC']>=1)   #...WHO WERE VACCINATED FOR CPOX

    had_cpox1= p_numvrc_M.where(p_numvrc_M['HAD_CPOX']==1)   # 1 : is for those who HAD CPOX
    had_cpox2= p_numvrc_M.where(p_numvrc_M['HAD_CPOX']==2)   # 2 : is for those who DIDN'T HAVE CPOX

    resultM= had_cpox1['SEX'].count()/had_cpox2['SEX'].count()   #Here we ensure that we take only the boys
    #print(resultM)

    res_dict= {"male":resultM.round(5),"female":resultF.round(5)}
    return res_dict    
    raise NotImplementedError()
print(chickenpox_by_sex())
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #               "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    df=pd.read_csv('assets/NISPUF17.csv',index_col=0)  #read the whole file
    df=df.loc[:,['HAD_CPOX', 'P_NUMVRC']]  #locate the specific columns of the csv file 
    df=df.loc[df['HAD_CPOX'].lt(3)].dropna()  #locate the 'HAD_CPOX' and leave out the 77 and 99 records
                                              # and drop all the NA of both the columns
    df
    corr, pval=stats.pearsonr(df['HAD_CPOX'],df['P_NUMVRC']) 
    return corr
    raise NotImplementedError()
print(corr_chickenpox())
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```